---
title: Basics
description: Basics
---

## Loading libraries

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

asdfasdfasdfasdfsadfasfasdfsasdffffffffffffffffffffffffffffffffffffffffffffffffffffff

`@instructions`
Load some library using library(X)

`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
